# Pyfar Open Educational Resources

Collection of open educational resources on acoustics.  
Visit [pyfar open educational resources](https://pyfar-oer.readthedocs.io/en/latest/) for more.

```python
import numpy as np
```

blablabla

<!-- grade_id: task, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c661359a14a8713d, points: 1, grade: true, solution: false -->
```python
assert x == 5
```