# Problem 1

![](pyfar_logos_fixed_size_pyfar.png)

blablabla, assignment blabla

```python
import numpy as np
```

```python
def do_somethin(a, b):
    return a*b
```

<!-- grade_id: answer1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test1, points: 3, grade: true, solution: false -->
```python
np.testing.assert_equal(c, 2)
```