# Problem 2

```python
import pyfar as pf
import numpy as np
```

```python
speech = pf.io.read_audio('impulse.wav')
```

Now do something cool

<!-- grade_id: answer1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test1, points: 2, grade: true, solution: false -->
```python
np.testing.assert_equal(something, 'cool')
```